---
title_meta  : Chapter 1 
title       : Introduction to networks
description : Basic concepts about ecological networks, 

--- type:VideoExercise lang:r xp:50 skills:1
## Introduction to ecological networks


*** =video_link
//player.vimeo.com/video/153476481

--- type:MultipleChoiceExercise lang:r xp:50 skills:1
## One way or another

If links in a network are represented as arrows, the network is:

*** =instructions
- Directed
- Undirected
- Mutualistic

*** =hint
Arrows usually show in which *direction* energy or biomass flows in the foodweb.

*** =pre_exercise_code
```{r}
# pec
```

*** =sct
```{r}
test_mc(1) # if 2 is the correct option.
```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1
## Win win relations

An ecological network encompassing relationships in wich both species benefit from them is a:

*** =instructions
- food web
- mutualistic network
- host-parasite network

*** =hint
In this type of network, relations have a mutual benefit. 

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
test_mc(2) # if 2 is the correct option.
```


--- type:NormalExercise lang:r xp:100 skills:1

<script type="text/javascript" src="http://latex.codecogs.com/latexit.js"></script>

## Network formats: edgelists and adjacency matrices

Several computer readable formats exist to store the description and information contained in a network. The most commonly used ones in ecological networks (and other disciplines) are: (i) adjacency matrices, and (ii) edgelists.

Adjancecy are matrices in which columns and rows represent the nodes/species in the network, and a link/interaction between nodes/species is repsented as a 1. The absence of an interacion is given by a 0.

Thus, a network can be defined by a matrix **A** where:

<div lang="latex">
\\begin{cases} 
      a_{ij} = 1 & if \ there \ is \ an \ interaction \ between \ i \ and \ j \\\\
      a_{ij} = 0 & if \ there \ is \ NO \ interaction \ between \ i \ and \ j
\\end{cases}
</div>


In the case of a foodweb, for example, as we saw previously, links are directed from prey to predator. The convention here is to take prey as being represented by rows and predators by columns. Thus, an interaction a<sub>ij</sub> runs from prey i to predator j.

In the case of a mutualistic (i.e. bipartite) network, species are not all represented in each dimension of the matrix (rows and columns). Instead hosts are usually shown in rows and their mutualistic partners are shown in columns. This makes these matrices non-square.

Edgelists on the other hand are lists of pairs of identifiers (usually integers) that
denote links (or ecological interactions) from the first item in pair to the second.
For example, a network specified by the following edgelist: 

1 2 <br>
2 3 <br>
1 3 <br>

is one with three nodes: V = {1,2,3}, and three links: one from basal prey 1 to primary consumer 2, another one from prey (primary consumer) 2 to top predator 3, and finally, an interaction between the basal resource and the top predator (E = {(1,2), (2,3), (1,3)}). This makes the top predator an omnivorous species. A representation of this network is shown below:

<p> <img src="https://raw.githubusercontent.com/mlurgi/images/master/net1.png" alt="Food chain"/>
</p>

This network can be represented as an adjacency matrix thus:

<p><img src="https://raw.githubusercontent.com/mlurgi/images/master/matrix.png" alt="Adjacency matrix"/>
</p>

More sophisticated ways of representing and storing networks exist, such as for example
the <a href="http://graphml.graphdrawing.org">GraphML</a> format, in which additional information about the network, the nodes, and 
the links can be included in the description.

In this course we will use edgelists as a starting format and then convert that 
representation into an adjacency matrix.


*** =instructions
- Make an adjacency matrix for the Benguela food web. For now, let it contain only zeros. 
- Add intereactions to the adjacency matrix.

*** =hint
- Remember that in numbers in an edge list represent species. To know how many species there are in the food web, look for the largest number in the edge list.
- Each line in the edge list represents an interaction between the species in each column.

*** =pre_exercise_code
```{r}
# pec
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

```

*** =sample_code
```{r}
# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- 

# Introduce ones to the matrix to represent interactions between species

```

*** =solution
```{r}
# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1
```

*** =sct
```{r}
test_error()

# Adjacency matrix
test_function(matrix, c("nrow", "ncol"))

# Interactions
i.m <- "It seems that you did not define the interactions correctly. Check your code and try again."
test_object(benguela.AM, undefined_msg = "Did you define an adjacency matrix called benguela.AM?",
            incorrect_msg = i.m)

success_msg("Fantastic job! Let's continue.")
```


--- type:MultipleChoiceExercise lang:r xp:75 skills:1
## Calculate number of nodes: food webs

In abstract representations of ecological communities as networks, nodes represent species and links amongst them denote ecological interactions happening between the represented species. 

Adjacency matrices are matrix representations of networks were elements in the matrix denote the presence of a link between the species represented by the row and the species represented by the column. To represent a food web composed of S species as a matrix then we need a SxS matrix in order to be able to represent all possible interaction between the S species. The number of species in the network can be calcuated thus as the number of rows or the number of columns.

The Benguela food web adjacency matrix `benguela.AM` that you created in the last exercise is already loaded. Calculate the number of species. <br>

*** =instructions
- 30
- 31
- 29
- 28

*** =hint
In an adjacency matrix, all species are represented by columns and rows. If you calculate the number of columns or rows of the matrix you will get the number of species in the food web. Also, you can use the `dim` function to look at the *dimensions* of the matrix.

*** =pre_exercise_code
```{r}
# pec
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1
```


*** =sct
```{r}
test_mc(3) 
```


--- type:MultipleChoiceExercise lang:r xp:75 skills:1
## Calculate number of links: food web

REMIND INTERACTIONS ARE REPRESENTED AS ONES Now calculate the number of links

*** =instructions
- 200
- 201
- 202
- 203
- 204

*** =hint
Since interactions are represented as ones in the matrix, if you `sum` all the values in the matrix, you will get the number of interactions.

*** =pre_exercise_code
```{r}
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1
```


*** =sct
```{r}
test_mc(4)
```


--- type:MultipleChoiceExercise lang:r xp:75 skills:1
## Calculate number of nodes: bipartite network

Food webs are network representation of who eats whom relationships. As such, there are not distinct subset of species among which interactions are not possible. This is why adjacency matrices for representing food webs contain the total number of species as rows and also as columns. 

Other types of ecological networks, however, like for example, plant-animal mutualistic networks or host-parasite interaction networks, exhibit a bipartite structure, in which species can be classified into two distinct groups where interactions do not occur among members of the same group. For example, in a plant-animal mutualistic network, plants do not interact with plants and animal mutualists do not interact among them. This feature allows for a more concise matrix representation of bipartite networks.

An adjacency matrix for a bipartite network only needs to represent one set of species (e.g. hosts/plants) as rows, and the other set of species (e.g. visitors/mutualists) as columns. Thus, the way of counting species in an adjacency matrix representing a mutualistic (bipartite) network is different from that used to count species in adjacency matrices representing food webs.

For an ecological network represented by an adjacency matrix with size HxV, where H is the number of hosts and V is the number of visitors, the number of species in the network is equal to H + V.

We have loaded a matrix representing a bipartite network of anemones with their interacting fish species called `anemonef`. Calculate the total number of anemone *and* fish species in the network.

*** =instructions
- 26
- 28
- 36
- 38
- 40

*** =hint
Remember that in a bipartite network, one set of species (in this case, anemones) are represented as columns, while the other set (in this case, fishes) are represented as rows.

*** =pre_exercise_code
```{r}
library(RCurl)
y <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/anemonefish.txt")
anemonef <- read.table(text = y)
names(anemonef) <- paste("A", 1:10, sep = "")
row.names(anemonef) <- paste("F", 1:26, sep = "")       
anemonef <- as.matrix(anemonef)
```


*** =sct
```{r}
test_mc(3)
```


--- type:NormalExercise lang:r xp:100 skills:1
## Weaving networks from scratch

EXPLAIN IGRAPH AND GIVE AN EXAMPLE OF HOW TO CREATE A SIMPLE NETWORK.

*** =instructions
- First create a vector called `species` with the numbers from 1 to 10
- Then create a network by adding the `species` as `vertices` to an empty graph.
- Create an interaction between species 5 and 7
- Create 10 other interactions between any species you want

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
require(igraph)
```

*** =sample_code
```{r}
# Ten species
species <- 
network <- 

# Link between species 5 and 7

```

*** =solution
```{r}
# Ten species
species <- 1:10
network <- graph.empty() + vertices(species)

# Link between species 5 and 7
network[5,7] <- 1
```

*** =sct
```{r}
# sct code
test_object(network, eq_condition = "identical",
            undefined_msg = "It seems you didn't create an object called <code>network</code>. Try again.",
            incorrect_msg = "Ups, it seems your network does not have an interaction between the right pair of species. Try again.")

success_msg("Perfect! That network looks just fine!")
```


--- type:NormalExercise lang:r xp:100 skills:1
## Converting matrix to igraph object

If we have an adjacency matrix and we want an igraph object representing the same network, without the need to create the network from scratch, igraph offers the functionality of creating graphs from adjancecy matrices. Just use the the function graph.adjacency(A), where A is your adjacency matrix.

Check out the documentation for that function <a href="http://igraph.org/r/doc/graph_from_adjacency_matrix.html">here</a>, including additional parameters that you can use to create your network!

*** =instructions
- Use the `graph.adjacency()` function to convert `benguela.AM` into an `igraph` network called `benguela.network`


*** =hint
The adjacency matrix `benguela.AM` is already loaded. Just apply `graph.adjacency()` to it.

*** =pre_exercise_code
```{r}
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1

require(igraph)
```

*** =sample_code
```{r}
# Convert Benguela adjacency matrix to an igraph network
benguela.network <- 
```

*** =solution
```{r}
# Convert Benguela adjacency matrix to an igraph network
benguela.network <- graph.adjacency(benguela.AM)
```

*** =sct
```{r}
test_error()
test_function(graph.adjacency)
success_msg("Well done! Now you can use all the power to igraph to analyse the Benguela food web")
```


--- type:NormalExercise lang:r xp:100 skills:1
## Calculate number of links and nodes with igraph

Once we have our network represented as an igraph object we can employ a wide range of functions provided by igraph to calculate properties of our network. Check out the igraph documentation for a full range of functions.

For example, the vcount() function from igraph allows you to count the number of nodes/species in your network/food web. Similarly, the ecount() function, lets you calculate the total number of links/interactions in your network.

*** =instructions
- Use function `vcount()` to count the number of vertices (i.e., species) in `benguela.network`
- Use function `ecount()` to count the number of edges (i.e., interactions) in `benguela.network`

*** =hint
The `benguela.network` is already loaded. You just have to use `vcount()` and `ecount()` on it.

*** =pre_exercise_code
```{r}
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1

require(igraph)
benguela.network <- graph.adjacency(benguela.AM)
```

*** =sample_code
```{r}
# Calculate number of species


# Calculate number of interactions

```

*** =solution
```{r}
# Calculate number of species
vcount(bengala.network)

# Calculate number of interactions
ecount(bengala.network)
```

*** =sct
```{r}
test_error()
test_function(vcount)
test_function(ecount)
success_msg("Good job!")
```


--- type:NormalExercise lang:r xp:100 skills:1
## Calculate connectance

*** =video_link
//player.vimeo.com/video/153476481

INSERT LINK FOR THIS VIDEO

Now that we know what connectance is and how to calculate it, let's use this knowledge to calculate the connectance of our networks.

For example, in the case of the adjacency matrix representing a food web we know that a matrix A contains all the links in our network. Hence, the number of links in our network is simply sum(A). Similarly we have also learnt that the number of species (at least for a food web) equals the number of rows (or columns) in the matrix: nrow(A).

The connectance of our food web represented by matrix A is thus: sum(A)/(nrow(A)**2)


*** =instructions
- Calculate connectance using igraph functions


*** =hint
Connectance is the number of interactions (`ecount()`), between the number of species squared (`vcount()^2`).

*** =pre_exercise_code
```{r}
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1

require(igraph)
benguela.network <- graph.adjacency(benguela.AM)
```

*** =sample_code
```{r}
# Calculate connectance
connectance <- 
        
# Print connectance value
print(connectance)
```

*** =solution
```{r}
# Calculate connectance
connectance <- ecount(bengala.network) / vcount(bengala.network)^2 

# Print connectance value
print(connectance)
```

*** =sct
```{r}
test_error()
test_object(connectance, "identical")
```


--- type:NormalExercise lang:r xp:100 skills:1
## Calculate links per species

Another way of assessing the connectivity of ecological networks is to calculate the mean number of links per species (L/S). The number of links per species offers a quick summary on how many interacting partners, on average, the species in the ecological network we are analysing have. Similar to connectance, it also gives us an idea of how connected are the species within our network.

*** =instructions
- Using igraph functions, calculate the average number of links per species in `benguela.network`


*** =hint
Just divide the number of links given by `ecount()` between the number of species given by `vcount()`.

*** =pre_exercise_code
```{r}
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1

require(igraph)
benguela.network <- graph.adjacency(benguela.AM)
```

*** =sample_code
```{r}
# Calculate number of links per species
linkes.per.species <- 
```

*** =solution
```{r}
# Calculate number of links per species
linkes.per.species <- ecount(beguela.network) / vcount(benguela.network)
```

*** =sct
```{r}
test_error()
test_object(links.per.species, "identical")
```


--- type:NormalExercise lang:r xp:200 skills:1
## Calculate connectance for bipartite network

We have learnt how to calculate the connectance of ecological networks in general: it is simply the fraction of realised links out of the possible ones. Additionally, we have explored methods to quantify connectance in food webs based on the information extracted either from the adjacency matrix or the igraph representation of our networks.

We have also learnt that the way in which bipartite networks are represented in adjacency matrices is different from the way in which we represent food webs. Keeping that in mind, we need to find out ways of finding out what would be: 1.- the total possible number of links in a bipartite network (remember that links between species in the same group are 'forbidden'), and 2.- the correct way of finding out the number of species in a bipartite network from our matrix representation.

*** =instructions
- Calculate connectance of the bipartite network `anemonef`


*** =hint
Remember that in bipartite networks, one set of species is represented as rows while the other is represented by columns.

*** =pre_exercise_code
```{r}
library(RCurl)
y <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/anemonefish.txt")
anemonef <- read.table(text = y)
names(anemonef) <- paste("A", 1:10, sep = "")
row.names(anemonef) <- paste("F", 1:26, sep = "")       
anemonef <- as.matrix(anemonef)
```

*** =sample_code
```{r}
# Calculate conectance of anemonef
bipartite.connectance <- 
        
```

*** =solution
```{r}
# Calculate conectance of anemoref
bipartite.connectance <- sum(anemonef) / dim(anemonef)[1] * dim(anemonef)[2]
        
```

*** =sct
```{r}
test_error()
test_object(bipartite.connectance, "identical")
```


--- type:MultipleChoiceExercise lang:r xp:100 skills:1
## The most connected

We have loaded 5 food webs (A, B, C, D, E). Which has the highest connectance?

*** =instructions
- A
- B
- C
- D
- E

*** =hint
Use what you have learned.  

*** =pre_exercise_code
```{r}
require(igraph)
set.seed(3)

A <- matrix(0, 15, 15)
A[matrix(sample(15, 76, replace = T), ncol = 2)] <- 1

B <- matrix(0, 52, 52)
B[matrix(sample(52, 800, replace = T), ncol = 2)] <- 1

C <- matrix(0, 151, 151)
C[matrix(sample(151, 1500, replace = T), ncol = 2)] <- 1

D <- matrix(0, 8, 8)
D[matrix(sample(8, 56, replace = T), ncol = 2)] <- 1

E <- matrix(0, 327, 327)
E[matrix(sample(327, 106, replace = T), ncol = 2)] <- 1
```

*** =sct
```{r}
test_mc(4) # if 2 is the correct option.
```


--- type:NormalExercise lang:r xp:100 skills:1
## Represent a network graphically

The network representations we are explored so far are great for manipulating them and calculating properties over them. However, to gain an intuitive idea of what the networks we are analysing look like, or if we want to show them in publications or websites, it is sometimes desirable to draw them as graphs: with dots representing nodes/species and lines connecting them representing links/interactions. In pretty much the same way we saw in the introductory video.

Of course, drawing a network on a whiteboard using coloured pens is easy. Doing that on a computer or writing code in R to do it, might be more of a challenge. Fortunately, libraries such as igraph offer these functionalities. The plot() command from the igraph library (<a href="http://igraph.org/r/doc/plot.igraph.html">documentation hereL</a>) takes as an input an igraph object and produces a graphical representation of the network.

Let's try it!

*** =instructions
- Plot Benguela food web


*** =hint
You can use the `plot()` function directly on `benguela.network`

*** =pre_exercise_code
```{r}
library(RCurl)
x <- getURL("https://raw.githubusercontent.com/seblun/networks_datacamp/master/datasets/benguela.tro")
benguela.EL <- read.table(text = x) 
benguela.EL <- as.matrix(benguela.EL)

# Create an adjacency matrix called benguela.AM, containing only zeros
benguela.AM <- matrix(0, max(benguela.EL), max(benguela.EL))

# Introduce ones to the matrix to represent interactions between species
benguela.AM[benguela.EL] <- 1

require(igraph)
benguela.network <- graph.adjacency(benguela.AM)
```

*** =sample_code
```{r}
# Plot the Benguela food web

```

*** =solution
```{r}
# Plot the Benguela food web
plot(benguela.network, edge.arrow.size = 0.2)
```

*** =sct
```{r}
test_error()
test_function(plot)
```